```python
version = "REPLACE_PACKAGE_VERSION"
```

---
# Assignment 1 Part 2: N-gram Language Models (Cont.) (30 pts)

In this assignment, we're going to train an n-gram language model that is able to "imitate" William Shakespeare's writing. 

```python
# Configure nltk

import nltk

nltk_data_path = "assets/nltk_data"
if nltk_data_path not in nltk.data.path:
    nltk.data.path.append(nltk_data_path)
```

<!-- grade_id: cell-125dc5645c8de34a, points: 0, grade: false, solution: true -->
```python
# Copy and paste the functions you wrote in Part 1 here and import any libraries necessary
# We have tried a more elegant solution by using
# from ipynb.fs.defs.assignment1_part1 import load_data, build_vocab, build_ngrams
# but it doesn't work with the autograder...

def load_data():
    """
    Load text data from a file and produce a list of token lists
    """
    import re
    import string
    from nltk.tokenize import sent_tokenize, word_tokenize

    text = open('assets/gutenberg/THE_SONNETS.txt','r').read().split('\n')

    text = [s.lower() for s in text]
    # text = [''.join(s for s in c if s not in string.punctuation) for c in text]
    text = [s.strip(' ') for s in text]
    text = [s for s in text if not s.isdigit()]
    text = [s for s in text if s]

    text = [sent_tokenize(s) for s in text]

    sentences = []

    for s in text:
        sentences.append(word_tokenize(s[0]))
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    return sentences

def build_vocab(sentences):
    """
    Take a list of sentences and return a vocab
    """
    
    #stu_ans = load_data()
    vocab = ['<s>', '</s>']
    for s in sentences:
        for c in s:
            vocab.append(c)

    vocab = list(set(vocab))

    # YOUR CODE HERE
    #raise NotImplementedError()
    
    return vocab

def build_ngrams(n, sentences):
    """
    Take a list of unpadded sentences and create all n-grams as specified by the argument "n" for each sentence
    """
    
    from nltk.util import ngrams

    all_ngrams = []

    for s in sentences:
        ngram = list(ngrams(s, n, pad_left=True, pad_right=True, left_pad_symbol='<s>', right_pad_symbol='</s>'))
        all_ngrams.append(ngram)

    
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    return all_ngrams
```

## Question 4: Guess the next token (20 pts)

With the help of the three functions you wrote in Part 1, let's first answer the following question as a review on $n$-grams.

Assume we are now working with bi-grams. What is the most likely token that comes after the sequence `<s> <s> <s>`, and how likely? Remember that a bi-gram language model is essentially a first-order Markov Chain. So, what determines the next state in a first-order Markov Chain? 

**Complete the function below to return a `tuple`, where `tuple[0]` is a `str` representing the mostly likely token and `tuple[1]` is a `float` representing its (conditional) probability of being the next token.**

<!-- grade_id: cell-f0590ac9b51f337e, points: 0, grade: false, solution: true -->
```python
def bigram_next_token(start_tokens=("<s>", ) * 3):
    """
    Take some starting tokens and produce the most likely token that follows under a bi-gram model
    """
    n = 2
    sentences = load_data()
    vocab = build_vocab(sentences)
    ngrams = build_ngrams(n, sentences)
    
    next_token, prob = None, None
    
    tokens = {}
    
    def flatten(t):
        return [item for sublist in t for item in sublist]
    
    ngrams_flat = flatten(ngrams)
    
    vocab_2 = flatten(ngrams)
    vocab_flat = flatten(vocab_2)
    
    
    for s in ngrams_flat:
        if s[0] == start_tokens[-1]:
            token = s[n-1]
            prob = (ngrams_flat.count(s))/vocab_flat.count(s[0])
                
#                 (vocab.count(start_tokens[-1])+1)/(len(vocab)*2) * 
#                     (ngrams_flat.count(s)+1)/(vocab.count(start_tokens[-1])+len(vocab))
#                    )
            tokens.update({token:prob})
        
    def sort_dict_by_value(d, reverse = True):
        return dict(sorted(d.items(), key = lambda x: x[1], reverse = reverse))
    
    sorted_dict = sort_dict_by_value(tokens)
    
    next_token = next(iter(sorted_dict))
    prob = sorted_dict[next_token]
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    return next_token, prob
    #return vocab_flat.count(s[0])
```

```python
# test = bigram_next_token()
# test



```

```python

```

```python
# n = 2
# sentences = load_data()
# vocab = build_vocab(sentences)
# ngrams = build_ngrams(n, sentences)

# next_token, prob = None, None

# tokens = {}





# def flatten(t):
#     return [item for sublist in t for item in sublist]

# flat = flatten(ngrams)
# flat
```

<!-- grade_id: cell-6a94a4513c67c4d5, points: 20, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = bigram_next_token(start_tokens=("<s>", ) * 3)

assert isinstance(stu_ans, tuple), "Q4: Your function should return a tuple. "
assert len(stu_ans) == 2, "Q4: Your tuple should have two elements. "
assert isinstance(stu_ans[0], str), "Q4: tuple[0] should be a str. "
assert isinstance(stu_ans[1], float), "Q4: tuple[1] should be a float. "

# Some hidden tests


del stu_ans
```

## Question 5: Train an $n$-gram language model (10 pts)

Now we are well positioned to start training an $n$-gram language model. We can fit a language model using the `MLE` class from `nltk.lm`. It requires two inputs: a list of all $n$-grams for each sentence and a vocabulary, both of which you have already written a function to build. Now it's time to put them together to work. 

**Complete the function below to return an `nltk.lm.MLE` object representing a trained $n$-gram language model.**

<!-- grade_id: cell-6136d38621377eeb, points: 0, grade: false, solution: true -->
```python
from nltk.lm import MLE

def train_ngram_lm(n):
    """
    Train a n-gram language model as specified by the argument "n"
    """
    sentences = load_data()
    vocab = build_vocab(sentences)
    ngrams = build_ngrams(n, sentences)
    
    lm = MLE(n)
    lm.fit(ngrams,vocab)
    
    
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    return lm
```

```python
# from nltk.lm import MLE
# n=2 

# sentences = load_data()
# vocab = build_vocab(sentences)
# ngrams = build_ngrams(n, sentences)

# lm = MLE(n)
# lm.fit(ngrams,vocab)




```

<!-- grade_id: cell-9c44953c467910db, points: 10, grade: true, solution: false -->
```python
# Autograder tests

stu_n = 4
stu_lm = train_ngram_lm(stu_n)
stu_vocab = build_vocab(load_data())

assert isinstance(stu_lm, nltk.lm.MLE), "Q3b: Your function should return an nltk.lm.MLE object. "

assert hasattr(stu_lm, "vocab") and len(stu_lm.vocab) == len(stu_vocab) + 1, "Q3b: Your language model wasn't trained properly. "

del stu_n, stu_lm, stu_vocab
```

FINALLY, are you ready to compose sonnets like the real Shakespeare?! We provide some starter code below, but absolutely feel free to modify any parts of it on your own. It'd be interesting to see how the "authenticity" of the sonnets is related to the parameter $n$. Do the sonnets feel more Shakespeare when you increase $n$? 

```python
# Every time it runs, depending on how drunk it is, a different sonnet is written. 
n = 10
num_lines = 14
num_words_per_line = 8
text_seed = ["<s>"] * (n - 1)

lm = train_ngram_lm(n)

sonnet = []
while len(sonnet) < num_lines:
    while True:  # keep generating a line until success
        try:
            line = lm.generate(num_words_per_line, text_seed=text_seed)
        except ValueError:  # the generation is not always successful. need to capture exceptions
            continue
        else:
            line = [x for x in line if x not in ["<s>", "</s>"]]
            sonnet.append(" ".join(line))
            break

# pretty-print your sonnet
print("\n".join(sonnet))
```

Output:
```
and to be praised of ages yet to
and though they be outstripped by every pen
but you shall shine more bright in these
or what strong hand can hold his swift
for that sweet odour , which doth in
when not to be , receives reproach of
my life being made of four , with
the teeming autumn big with rich increase ,
and both for my sake lay on me
when beauty lived and died as flowers do
and you but one , can every shadow
my saucy bark ( inferior far to his
or else of thee this i prognosticate ,
o thou my lovely boy who in thy
```